# Data structures

Once we start working with many pieces of data at once, it will be convenient for us to store data in structures like arrays or dictionaries (rather than just relying on variables).

Types of data structures covered:
1. Tuples
2. Dictionaries
3. Arrays


As an overview, tuples and arrays are both ordered sequences of elements (so we can index into them). 

Dictionaries and arrays are both **mutable**, whereas tuples are **immutable**

We'll explain this more below!
Arrays (in my opinion) are one of the most important structures in (Scientific) Computing.

## Tuples

We can create a tuple by enclosing an ordered collection of elements in `( )`.

Syntax: <br>
```julia
(item1, item2, ...)
```


```julia
myfavoriteanimals = ("penguins", "cats", "sugargliders")
```

```julia
numbers = (1, 1, 2)
```

We can index into this tuple,

```julia
myfavoriteanimals[1]
```

but since tuples are immutable, we can't update it

```julia
myfavoriteanimals[1] = "otters"
```

```julia
numbers[1] = 5
```

Yes, Julia is 1-based indexing, not 0-based like Python.  (Although https://docs.julialang.org/en/v1/devdocs/offset-arrays/)

## Julia Specific: NamedTuples

As you might guess, `NamedTuple`s are just like `Tuple`s except that each element additionally has a name! They have a special syntax using `=` inside a tuple:

```julia
(name1 = item1, name2 = item2, ...)
```

```julia
myfavoriteanimals = (bird = "penguins", mammal = "cats", marsupial = "sugargliders")
```

Like regular `Tuples`, `NamedTuples` are ordered, so we can retrieve their elements via indexing:

```julia
myfavoriteanimals[1]
```

They also add the special ability to access values by their name:

```julia
myfavoriteanimals.bird
```

## Dictionaries

If we have sets of data related to one another, we may choose to store that data in a dictionary. We can create a dictionary using the `Dict()` function, which we can initialize as an empty dictionary or one storing key, value pairs.

Syntax:
```julia
Dict(key1 => value1, key2 => value2, ...)
```

A good example is a contacts list, where we associate names with phone numbers.

```julia
phonebook = Dict("Yiannis" => "0044000000", "Imperial" => "02075895111")
```

Output:
```
Dict{String, String} with 2 entries:
  "Yiannis"  => "0044000000"
  "Imperial" => "02075895111"
```

In this example, each name and number is a "key" and "value" pair. We can grab Yianni's number (a value) using the associated key

```julia
phonebook["Yiannis"]
```

We can add another entry to this dictionary as follows

```julia
phonebook["Graduate School"] = "020 7589 5111"
```

Let's check what our phonebook looks like now...

```julia
phonebook
```

We can delete Yiannis from our contact list - and simultaneously grab his number - by using `pop!`

```julia
yiannis_number = pop!(phonebook, "Yiannis")
```

Output:
```
"0044000000"
```

```julia
phonebook
```

Output:
```
Dict{String, String} with 1 entry:
  "Imperial" => "02075895111"
```

Unlike tuples and arrays, dictionaries are not ordered. So, we can't index into them.

```julia
phonebook[1]
```

In the cell above, `julia` thinks you're trying to access a value associated with the key `1`.

## Arrays

Unlike tuples, arrays are mutable. Unlike dictionaries, arrays contain ordered collections. <br>
We can create an array by enclosing this collection in `[ ]`.

Syntax: <br>
```julia
[item1, item2, ...]
```

For an interesting piece of Julia history regarding differences between matrices and vectors feel free to read https://github.com/julialang/julia/issues/4774


For example, we might create an array to keep track of pokemon


```julia
pokemon = ["Bulbasaur", "Squirtle", "Charmander", "Pikachu"]
```

Output:
```
4-element Vector{String}:
 "Bulbasaur"
 "Squirtle"
 "Charmander"
 "Pikachu"
```

```julia
typeof(pokemon)
```

Output:
```
Stacktrace:
 [1] top-level scope
   @ In[17]:1
```

`Vector` here is shorthand for a 1D array. The `{String}` following it indicates the elements are of the type `String`. It is an alias for `Array{String, 1}` - check it out:

We can store a sequence of numbers, or mixed type sequences

```julia
fibonacci = [1, 1, 2, 3, 5, 8, 13]
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

```julia
mixture = [1, 1, 2, 3, "Yiannis", "Jay"]
```

Output:
```
6-element Vector{Any}:
 1
 1
 2
 3
  "Yiannis"
  "Jay"
```

Once we have an array, we can grab individual pieces of data from inside that array by indexing into the array. For example, if we want the third pokemon listed in `pokemon`, we write

```julia
pokemon[3]
```

Output:
```
"Charmander"
```

We can use indexing to edit an existing element of an array

```julia
pokemon[3] = "Mewtwo"
```

Output:
```
"Mewtwo"
```

We can also edit the array by using the `push!` and `pop!` functions. `push!` adds an element to the end of an array and `pop!` removes the last element of an array. (also see `pushfirst!` and `popfirst!`)

Remember the Julia notation `!` means that the object changes.

We can add another number to our fibonnaci sequence

```julia
push!(fibonacci, 21)
```

and then remove it

```julia
pop!(fibonacci)
```

```julia
fibonacci
```

So far we've given examples of only 1D arrays of scalars, but arrays can have an arbitrary number of dimensions and can also store other arrays. 
<br><br>
For example, the following are arrays of arrays:

```julia
favorites = [["ice-cream", "chocolate", "eggs"],["penguins", "cats", "spiders"]]
```

Output:
```
2-element Vector{Vector{String}}:
 ["ice-cream", "chocolate", "eggs"]
 ["penguins", "cats", "spiders"]
```

```julia
numbers = [[1, 2, 3], [4, 5], [6, 7, 8, 9]]
```

Output:
```
3-element Vector{Vector{Int64}}:
 [1, 2, 3]
 [4, 5]
 [6, 7, 8, 9]
```

Below are examples of 2D and 3D arrays populated with random values.

```julia
rand(4, 3) # Matrix is shorthand for a 2D array!
```

Output:
```
4×3 Matrix{Float64}:
 0.286473  0.348052  0.824799
 0.438924  0.93903   0.0892267
 0.360405  0.646196  0.948074
 0.822157  0.998186  0.995855
```

```julia
rand(2, 4, 3)
```

Output:
```
2×4×3 Array{Float64, 3}:
[:, :, 1] =
 0.116875  0.769756  0.451757  0.766136
 0.216607  0.85258   0.459027  0.823307

[:, :, 2] =
 0.736376  0.788956  0.654249  0.218452
 0.662208  0.532137  0.460203  0.259198

[:, :, 3] =
 0.78053   0.879242  0.0129466  0.325102
 0.281126  0.50319   0.20404    0.191328
```

```julia
?rand
```

### Mutability

Be careful when you want to copy arrays!

```julia
fibonacci
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

```julia
fib = fibonacci
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

```julia
fib[1] = 404
```

Output:
```
404
```

```julia
fibonacci
```

Output:
```
7-element Vector{Int64}:
 404
   1
   2
   3
   5
   8
  13
```

Editing `fib` caused `fibonacci` to get updated as well!

In the above example, we didn't actually make a copy of `fibonacci`. We just created a new way to access the entries in the array bound to `fibonacci`.

If we'd like to make a copy of the array bound to `fibonacci`, we can use the `copy` function.

```julia
# First, restore fibonacci
fibonacci[1] = 1
fibonacci
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

```julia
fib2 = copy(fibonacci)
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

```julia
fib2[1] = 404
```

Output:
```
404
```

```julia
fibonacci
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

In this last example, fibonacci was not updated. Therefore we see that the arrays bound to `fib2` and `fibonacci` are distinct.

# Slicing

Let's be more creative with our indexing

```julia
squares = [1, 4, 9, 16, 25, 36, 49, 64, 81]
```

Output:
```
9-element Vector{Int64}:
  1
  4
  9
 16
 25
 36
 49
 64
 81
```

```julia
squares[1:3]
```

Output:
```
3-element Vector{Int64}:
 1
 4
 9
```

```julia
squares[1:2:end]
```

Output:
```
5-element Vector{Int64}:
  1
  9
 25
 49
 81
```

```julia
squares[end]
```

Output:
```
81
```

## Concatenation
If, instead of commas, we use spaces, then the values are concatenated horizontally

```julia
cubes = [1, 8, 27, 64, 125, 216, 343, 512]
```

Output:
```
8-element Vector{Int64}:
   1
   8
  27
  64
 125
 216
 343
 512
```

```julia
cubes2 = [1 8 27]
```

Output:
```
1×3 Matrix{Int64}:
 1  8  27
```

```julia
pop!(squares)
```

Output:
```
81
```

```julia
powers = [1:8 squares cubes]
```

Output:
```
8×3 Matrix{Int64}:
 1   1    1
 2   4    8
 3   9   27
 4  16   64
 5  25  125
 6  36  216
 7  49  343
 8  64  512
```

```julia
powers[4, 2]
```

Output:
```
16
```

```julia
powers[:, 3]
```

Output:
```
8-element Vector{Int64}:
   1
   8
  27
  64
 125
 216
 343
 512
```

```julia
powers[7, :]
```

Output:
```
3-element Vector{Int64}:
   7
  49
 343
```

Semicolon separators perform vertical concatenation:

```julia
[squares; cubes]
```

Output:
```
17-element Vector{Int64}:
   1
   4
   9
  16
  25
  36
  49
  64
  81
   1
   8
  27
  64
 125
 216
 343
 512
```

Whereas commas would simply create an array of arrays : 

```julia
nested_powers = [[1,2,3,4,5,6,7,8], squares, cubes]
```

Output:
```
3-element Vector{Vector{Int64}}:
 [1, 2, 3, 4, 5, 6, 7, 8]
 [1, 4, 9, 16, 25, 36, 49, 64, 81]
 [1, 8, 27, 64, 125, 216, 343, 512]
```

```julia
nested_powers[2]
```

Output:
```
9-element Vector{Int64}:
  1
  4
  9
 16
 25
 36
 49
 64
 81
```

Horizontal and vertical concatenation can be used together to as a simple syntax for matrix literals

```julia
[1 3 5; 2 4 6]
```

Output:
```
2×3 Matrix{Int64}:
 1  3  5
 2  4  6
```

### Exercises

#### 4.1 
Try to add "Emergency" as key to `phonebook` with the value `string(911)` with the following code
```julia
phonebook["Emergency"] = 911
```

Why doesn't this work?

```julia
phonebook["Emergency"] = 911 # wrong type! "911" will work
```

Output:
```
"911"
```

#### 4.2
Create a new dictionary called `flexible_phonebook` that has Jenny's number (**9671111**) stored as an integer and the Ghostbusters' number (**1-800-GHOSTS**) stored as a string with the following code

```julia
flexible_phonebook = Dict("Jenny" => 9671111, "Ghostbusters" => "1-800-GHOSTS")
```

Output:
```
Dict{String, Any} with 2 entries:
  "Jenny"        => 9671111
  "Ghostbusters" => "1-800-GHOSTS"
```

<!-- grade_id: cell-66f1d636c676c2f2, points: 1, grade: true, solution: false -->
```julia
@assert Pair("Jenny", 9671111) in flexible_phonebook
@assert Pair("Ghostbusters", "1-800-GHOSTS") in flexible_phonebook
```