# Loops & Comprehensions

## for loops

The syntax for a `for` loop is

```julia
for *var* in *loop iterable*
    *loop body*
end
```

## while loops

The syntax for a `while` is

```julia
while *condition*
    *loop body*
end
```



What do we think the following will produce?

```julia
i = 1
while i <= 10
    println(i)
    i += 1
end
```

Output:
```
1
2
3
4
5
6
7
8
9
10
```

`for` loops are better suited to these kinds of iteration. Write the equivalent `for` loop below

```julia
for i in 1:10 # Careful of 1-based indexing when working with arrays...
    println(i)
end
```

Output:
```
1
2
3
4
5
6
7
8
9
10
```

We could use a for loop to generate the same powers matrix as in the example above:

```julia
A = fill(0, (8, 3)) # Allocate an 8x3 matrix to store the values into
for pow in 1:3
    for value in 1:8
        A[value, pow] = value ^ pow
    end
end
A
```

Output:
```
8×3 Matrix{Int64}:
 1   1    1
 2   4    8
 3   9   27
 4  16   64
 5  25  125
 6  36  216
 7  49  343
 8  64  512
```

```julia
A == powers
```

Output:
```
true
```


## Array Comprehensions

```julia
squares = [value^2 for value in 1:8]
```

Output:
```
8-element Vector{Int64}:
  1
  4
  9
 16
 25
 36
 49
 64
```

```julia
cubes = [value^3 for value in 1:8]
```

Output:
```
8-element Vector{Int64}:
   1
   8
  27
  64
 125
 216
 343
 512
```

```julia
powers = [value^pow for value in 1:8, pow in 1:3]
```

Output:
```
8×3 Matrix{Int64}:
 1   1    1
 2   4    8
 3   9   27
 4  16   64
 5  25  125
 6  36  216
 7  49  343
 8  64  512
```

# The element type

Note that every time an array prints out, it is displaying its element type and dimensionality, for example `Array{Int64, 2}`. This describes what it can store — and thus what it can return upon indexing.

```julia
typeof(powers)
```

Output:
```
Matrix{Int64}[90m (alias for [39m[90mArray{Int64, 2}[39m[90m)[39m
```

```julia
typeof(powers[1, 1])
```

Output:
```
Int64
```

Further, the array will *try* to convert any new values assigned into it to its element type:

```julia
powers[1, 1] = -5.0 # This can be losslessly converted to an integer
```

Output:
```
-5.0
```

```julia
powers[1, 1] = 1.6 # TODO: fix this error
```

Output:
```
@ Base ./array.jl:968
 [4] top-level scope
   @ In[122]:1
```

```julia
powers
```

Output:
```
8×3 Matrix{Int64}:
 -5   1    1
  2   4    8
  3   9   27
  4  16   64
  5  25  125
  6  36  216
  7  49  343
  8  64  512
```

Arrays that have an exact and concrete element type are generally significantly faster, so Julia will try to find an amenable element type for you in its literal construction syntax:

```julia
fortytwosarray = [42, 42.0, 4.20e1, 4.20f1, 84//2, 0x2a]
```

Output:
```
6-element Vector{Float64}:
 42.0
 42.0
 42.0
 42.0
 42.0
 42.0
```

```julia
for x in fortytwosarray
    show(x)
    println("\tisa $(typeof(x))")
end
```

Output:
```
42	isa Int64
42.0	isa Float64
42.0	isa Float64
42.0f0	isa Float32
42//1	isa Rational{Int64}
0x2a	isa UInt8
```

The `Any` array can be helpful for disabling these behaviors and allowing all kinds of different objects:

```julia
fortytwosarray = Any[42, 42.0, 4.20e1, 4.20f1, 84//2, 0x2a]
```

Output:
```
6-element Vector{Any}:
   42
   42.0
   42.0
   42.0f0
  42//1
 0x2a
```

```julia
fortytwosarray[1] = "FORTY TWO"
fortytwosarray
```

Output:
```
@ Base ./array.jl:966
 [2] top-level scope
   @ In[137]:1
```

### Exercises


#### 5.1

Use a for loop to create the below matrix, called `a_matrix`:
```
3×3 Matrix{Int64}:
 1  2  3
 4  5  6
 7  8  9
```

```julia
a_matrix = fill(0, (3, 3))
for i in 1:3
    for j in 1:3
        a_matrix[j, i] = i + 3 * (j - 1)
    end
end
a_matrix
```

Output:
```
3×3 Matrix{Int64}:
 1  2  3
 4  5  6
 7  8  9
```

<!-- grade_id: cell-a6b20b78fbfeb449, points: 1, grade: true, solution: false -->
```julia
@assert a_matrix == [1 2 3 ; 4 5 6 ; 7 8 9]
```

#### 5.2
Use an array comprehension to create an an array `squares_arr` that stores the squares for all integers between 1 and 100.

```julia
squares_arr = [x^2 for x in 1:100]
```

Output:
```
100-element Vector{Int64}:
     1
     4
     9
    16
    25
    36
    49
    64
    81
   100
   121
   144
   169
     ⋮
  7921
  8100
  8281
  8464
  8649
  8836
  9025
  9216
  9409
  9604
  9801
 10000
```

```julia
@assert length(squares_arr) == 100
@assert sum(squares_arr) == 338350
```