## Introduction
Hello everyone, and welcome to the Introduction to Julia course. While Julia does have a command line interface, we will be using Jupyter (Julia, Python, R) notebooks. Incase you haven't used these before and to familiarize ourselves with them, we will be going over some basic programming concepts before moving on to more advanced topics.

# Getting started

Topics:
1.  How to print
2. How to assign variables
3. How to comment
4. Syntax for basic math

## How to print

In Julia we usually use `println()` to print

```julia
println("I'm excited to learn Julia!")
```

Output:
```
I'm excited to learn Julia!
```

And to access documentation for the println() and other functions


```julia
?println
```

Output:
```
search: [0m[1mp[22m[0m[1mr[22m[0m[1mi[22m[0m[1mn[22m[0m[1mt[22m[0m[1ml[22m[0m[1mn[22m [0m[1mp[22m[0m[1mr[22m[0m[1mi[22m[0m[1mn[22m[0m[1mt[22msty[0m[1ml[22med [0m[1mp[22m[0m[1mr[22m[0m[1mi[22m[0m[1mn[22m[0m[1mt[22m s[0m[1mp[22m[0m[1mr[22m[0m[1mi[22m[0m[1mn[22m[0m[1mt[22m is[0m[1mp[22m[0m[1mr[22m[0m[1mi[22m[0m[1mn[22m[0m[1mt[22m

[36m  println([io::IO], xs...)[39m

  Print (using [36mprint[39m) [36mxs[39m to [36mio[39m followed by a newline. If [36mio[39m is not supplied,
  prints to the default output stream [36mstdout[39m.

  See also [36mprintstyled[39m to add colors etc.

[1m  Examples[22m
[1m  ≡≡≡≡≡≡≡≡≡≡[22m

[36m  julia> println("Hello, world")[39m
[36m  Hello, world[39m
[36m  [39m
[36m  julia> io = IOBuffer();[39m
[36m  [39m
[36m  julia> println(io, "Hello", ',', " world.")[39m
[36m  [39m
[36m  julia> String(take!(io))[39m
[36m  "Hello, world.\n"[39m
```

## How to assign variables

All we need is a variable name, value, and an equal's sign!<br>
Julia is a dynamic language and thus will figure out types for us.

```julia
my_answer = 42
typeof(my_answer)
```

Output:
```
Int64
```

```julia
my_pi = 3.14159
typeof(my_pi)
```

Output:
```
Float64
```

```julia
pi
```

Output:
```
π = 3.1415926535897...
```

```julia
typeof(pi)
```

Output:
```
Irrational{:π}
```

```julia
😺 = "smiley cat!"
typeof(😺)
```

Output:
```
String
```

```julia

```

To type a smiley cat, use tab completion to select the emoji name and then tab again

```julia
# \:smi + <tab> --> select with down arrow + <enter> ---> <tab> + <enter> to complete
```

After assigning a value to a variable, we can reassign a value of a different type to that variable without any issue.

```julia
😺 = 1
```

Output:
```
1
```

```julia
typeof(😺)

```

Output:
```
Int64
```

## How to comment

```julia
# You can leave comments on a single line using the pound/hash key
```

```julia
#=
this is a further comment
For multi-line comments, 
use the '#= =#' sequence.

=#
```

## Syntax for basic math

```julia
sum = 3 + 7
```

Output:
```
10
```

```julia
difference = 10 - 3
```

Output:
```
7
```

```julia
product = 20 * 5
```

Output:
```
100
```

```julia
quotient = 100 / 10
```

Output:
```
10.0
```

```julia
power = 10 ^ 2
```

Output:
```
100
```

```julia
modulus = 101 % 2
```

Output:
```
1
```

# Numbers
Many different ways to write the number forty-two


```julia
fortytwos = (42, 42.0, 4.20e1, 4.20f1, 84//2, 0x2a)
```

Output:
```
(42, 42.0, 42.0, 42.0f0, 42//1, 0x2a)
```

```julia
for x in fortytwos
    println("$x \t isa \t $(typeof(x))")
end
```

Output:
```
42 	 isa 	 Int64
42.0 	 isa 	 Float64
42.0 	 isa 	 Float64
42.0 	 isa 	 Float32
42//1 	 isa 	 Rational{Int64}
42 	 isa 	 UInt8
```

# Updating Operators
All the infix arithmetic operators above can be used as "updating" operators in conjunction with an assignment:

```julia
y = 1
y += 1

```

Output:
```
2
```

```julia
println(y)

```

Output:
```
2
```

```julia
y = 1
y = y + 1

```

Output:
```
2
```

```julia
println(y)
```

Output:
```
2
```

Note that it's just re-purposing the same name for a new value. This means that the type can change!

```julia
y /= 2
```

Output:
```
1.0
```

```julia
typeof(y)
```

Output:
```
Float64
```

# Strings and characters

```julia
s1 = "hello world"

```

Output:
```
"hello world"
```

```julia
s2 = """ hello world"""

```

Output:
```
" hello world"
```

```julia
s3 = "Here, we get an "error" because it's ambiguous where this string ends "
```

Output:
```
Stacktrace:
 [1] top-level scope
   @ In[36]:1
```

```julia
s4 = """ no "errors" here """
```

Output:
```
" no \"errors\" here "
```

```julia
println("""The other nice thing about triple-quoted
               string literals is that they ignore leading
               indentation, which is nice for long strings
               in real code. Try changing these quotes!""")


```

Output:
```
The other nice thing about triple-quoted
string literals is that they ignore leading
indentation, which is nice for long strings
in real code. Try changing these quotes!
```

Strings are not written with single 's — that's used for a single character:

```julia
first(s1)
```

Output:
```
'h': ASCII/Unicode U+0068 (category Ll: Letter, lowercase)
```

```julia
c1 = 's'
```

Output:
```
's': ASCII/Unicode U+0073 (category Ll: Letter, lowercase)
```

```julia
'If you try writing a string in single-quotes, you will get an error'
```

# String interpolation


You can use the dollar sign inside a string to evaluate a Julia expression inside a string — either a single variable or a more complicated expression:


```julia
name = "Yiannis"
num_fingers = 10
num_toes = 10 
println("Hello, my name is $name.")
println("I have $num_fingers fingers and $num_toes toes.")
```

Output:
```
Hello, my name is Yiannis.
I have 10 fingers and 10 toes.
```

What do you think the following will print?

```julia
println("That is $(num_fingers+num_toes) digits in all")
```

Output:
```
That is 20 digits in all
```

### Exercises

#### 1.1 
Look up docs for the `convert` function.

```julia
?convert
```

Output:
```
search: [0m[1mc[22m[0m[1mo[22m[0m[1mn[22m[0m[1mv[22m[0m[1me[22m[0m[1mr[22m[0m[1mt[22m [0m[1mC[22m[0m[1mo[22m[0m[1mn[22mcurrency[0m[1mV[22miolation[0m[1mE[22m[0m[1mr[22mror [0m[1mc[22m[0m[1mo[22mde_[0m[1mn[22mati[0m[1mv[22m[0m[1me[22m @[0m[1mc[22m[0m[1mo[22mde_[0m[1mn[22mati[0m[1mv[22m[0m[1me[22m

[36m  convert(T, x)[39m

  Convert [36mx[39m to a value of type [36mT[39m.

  If [36mT[39m is an [36mInteger[39m type, an [36mInexactError[39m will be raised if [36mx[39m is not
  representable by [36mT[39m, for example if [36mx[39m is not integer-valued, or is outside
  the range supported by [36mT[39m.

[1m  Examples[22m
[1m  ≡≡≡≡≡≡≡≡≡≡[22m

[36m  julia> convert(Int, 3.0)[39m
[36m  3[39m
[36m  [39m
[36m  julia> convert(Int, 3.5)[39m
[36m  ERROR: InexactError: Int64(3.5)[39m
[36m  Stacktrace:[39m
[36m  [...][39m

  If [36mT[39m is a [36mAbstractFloat[39m type, then it will return the
```

#### 1.2 
Assign `365` to a variable named `days`. Convert `days` to a float and assign it to variable `days_float`
If you've done this correctly, the @assert function will check it for you.

```julia
days = 365
days_float = convert(Int16, days)
```

Output:
```
365
```

<!-- grade_id: cell-715f78016beb0489, points: 1, grade: true, solution: false -->
```julia
@assert days == 365
@assert days_float == 364

```

Output:
```
Stacktrace:
 [1] top-level scope
   @ In[87]:2
```

#### 1.3 
See what happens when you execute

```julia
convert(Int64, "1")
```
and

```julia
parse(Int64, "1")
```

```julia
convert(Int64, "1")
```

```julia

```

```julia

```

```julia

```