## Introduction
Hello everyone, and welcome to the Introduction to Julia course. While Julia does have a command line interface, we will be using Jupyter (Julia, Python, R) notebooks. Incase you haven't used these before and to familiarize ourselves with them, we will be going over some basic programming concepts before moving on to more advanced topics.

# Getting started

Topics:
1.  How to print
2. How to assign variables
3. How to comment
4. Syntax for basic math

## How to print

In Julia we usually use `println()` to print

```julia
println("Hello, World!")
```

Output:
```
Hello, World!
```

And to access documentation for the println() and other functions


```julia

```

Output:
```
Stacktrace:

 [1] top-level scope

   @ ~/ic-teach/RCDS-introduction-to-julia/day1/01. Getting started.ipynb:1
```

## How to assign variables

All we need is a variable name, value, and an equal's sign!<br>
Julia is a dynamic language and thus will figure out types for us.

```julia

```

```julia

```

```julia

```

```julia

```

```julia
😺 = "smiley cat!"
typeof(😺)
```

To type emojis and other characters, use tab completion to select the emoji name and then tab again

```julia
# \:smi + <tab> --> select with down arrow + <enter> ---> <tab> + <enter> to complete
# \beta + <tab>
```

```julia

```

After assigning a value to a variable, we can reassign a value of a different type to that variable without any issue.

```julia
😺 = 1
```

```julia
typeof(😺)

```

## How to comment

```julia

```

```julia

```

## Syntax for basic math

```julia

```

```julia

```

```julia

```

```julia

```

```julia

```

```julia

```

# Numbers
Many different ways to write the number forty-two


```julia

```

```julia
for x in fortytwos
    println("$x \t isa \t $(typeof(x))")
end
```

# Updating Operators
All the infix arithmetic operators above can be used as "updating" operators in conjunction with an assignment:

```julia

```

```julia
println(y)

```

```julia

```

```julia
println(y)
```

Note that it's just re-purposing the same name for a new value. This means that the type can change!

```julia

```

```julia
typeof(y)
```

# Strings and characters

```julia

```

```julia

```

```julia

```

```julia

```

```julia

```

Strings are not written with single 's — that's used for a single character:

```julia

```

```julia

```

```julia

```

# String interpolation


You can use the dollar sign inside a string to evaluate a Julia expression inside a string — either a single variable or a more complicated expression:


```julia
name = "Yiannis"
num_fingers = 10
num_toes = 10 
```

What do you think the following will print?

```julia

```

### Exercises

#### 1.1 
Look up docs for the `convert` function.

```julia
## YOUR CODE GOES HERE
```

#### 1.2 
Assign `365` to a variable named `days`. Convert `days` to a float and assign it to variable `days_float`
If you've done this correctly, the @assert function will check it for you.

```julia
## YOUR CODE GOES HERE
```

<!-- grade_id: cell-715f78016beb0489, points: 1, grade: true, solution: false -->
```julia
@assert days == 365
@assert days_float == 364

```

#### 1.3 
See what happens when you execute

```julia
convert(Int64, "1")
```
and

```julia
parse(Int64, "1")
```

```julia
## YOUR CODE GOES HERE
```

```julia
## YOUR CODE GOES HERE
```

```julia
# Look up the documentation for parse
```