# Loops & Comprehensions

## for loops

The syntax for a `for` loop is

```julia
for *var* in *loop iterable*
    *loop body*
end
```

## while loops

The syntax for a `while` is

```julia
while *condition*
    *loop body*
end
```



What do we think the following will produce?

```julia
i = 1
while i <= 10
    println(i)
    i += 1
end
```

`for` loops are better suited to these kinds of iteration. Write the equivalent `for` loop below

```julia
## YOUR CODE GOES HERE
```

We could use a for loop to generate the same powers matrix as in the example above:

```julia
A = fill(0, (8, 3)) # Allocate an 8x3 matrix to store the values into
for pow in 1:3
    for value in 1:8
        A[value, pow] = value ^ pow
    end
end
A
```

```julia
A == powers
```


## Array Comprehensions

```julia

```

```julia

```

```julia

```

# The element type

Note that every time an array prints out, it is displaying its element type and dimensionality, for example `Array{Int64, 2}`. This describes what it can store — and thus what it can return upon indexing.

```julia
typeof(powers)
```

```julia
typeof(powers[1, 1])
```

Further, the array will *try* to convert any new values assigned into it to its element type:

```julia
powers[1, 1] = -5.0 # This can be losslessly converted to an integer
```

```julia
powers[1, 1] = 1.6 # TODO: fix this error
```

```julia
powers
```

Arrays that have an exact and concrete element type are generally significantly faster, so Julia will try to find an amenable element type for you in its literal construction syntax:

```julia
fortytwosarray = [42, 42.0, 4.20e1, 4.20f1, 84//2, 0x2a]
```

```julia
for x in fortytwosarray
    show(x)
    println("\tisa $(typeof(x))")
end
```

The `Any` array can be helpful for disabling these behaviors and allowing all kinds of different objects:

```julia
fortytwosarray = Any[42, 42.0, 4.20e1, 4.20f1, 84//2, 0x2a]
```

```julia
fortytwosarray[1] = "FORTY TWO"
fortytwosarray
```

### Exercises


#### 5.1

Use a for loop to create the below matrix, called `a_matrix`:
```
3×3 Matrix{Int64}:
 1  2  3
 4  5  6
 7  8  9
```

```julia
## YOUR CODE GOES HERE
```

<!-- grade_id: cell-a6b20b78fbfeb449, points: 1, grade: true, solution: false -->
```julia
@assert a_matrix == [1 2 3 ; 4 5 6 ; 7 8 9]
```

#### 5.2
Use an array comprehension to create an an array `squares_arr` that stores the squares for all integers between 1 and 100.

```julia
## YOUR CODE GOES HERE
```

```julia
@assert length(squares_arr) == 100
@assert sum(squares_arr) == 338350
```