# Basic linear algebra in Julia
Based on work by Andreas Noack Jensen (MIT) (https://github.com/andreasnoack)


First let's define a random matrix

```julia
A = rand(1:4, 3, 3)
```

Define a vector of ones

```julia
x = ones(3)
```

Many of the basic operations are the same as in other languages

```julia
b = A * x
```


#### Transposition

```julia
A'
```

and we can get the transpose with

```julia
transpose(A)
```

#### Transposed multiplication
Julia allows us to write this without *

```julia
A'A
```

#### Solving linear systems 
The problem $Ax=b$ for ***square*** $A$ is solved by the \ operator.

```julia
x = A \ b
```

`A\b` gives us the *least squares solution* if we have an overdetermined linear system (a "tall" matrix)

```julia
Atall = rand(3, 2)
```

```julia
Atall \ b
```

and the *minimum norm least squares solution* if we have a rank-deficient least squares problem

```julia
v = rand(3)
rankdef = hcat(v, v)
```

```julia
rankdef \ b
```

Julia also gives us the minimum norm solution when we have an underdetermined solution (a "short" matrix)

```julia
bshort = rand(2)
Ashort = rand(2, 3)
```

```julia
Ashort \ bshort
```

# The LinearAlgebra library

While much of linear algebra is available in Julia by default (as shown above), there's a standard library named `LinearAlgebra` that brings in many more relevant names and functions. In particular, it provides factorizations and some structured matrix types.  As with all packages, you can bring these additional features into your session with a `using LinearAlgebra`.

### Exercises

#### 10.1 
Take the inner product (or "dot" product) of a vector `v` with itself and assign it to variable `dot_v`. 



```julia
## YOUR CODE GOES HERE
```

```julia
## YOUR CODE GOES HERE
```

```julia
## YOUR CODE GOES HERE
```

<!-- grade_id: cell-913fef9b0d19cd52, points: 1, grade: true, solution: false -->
```julia
@assert dot_v == 14
```

#### 10.2 
Take the outer product of a vector v with itself and assign it to variable `outer_v`

```julia
## YOUR CODE GOES HERE
```

```julia
@assert outer_v == [1 2 3
                    2 4 6
                    3 6 9]
```

#### 10.3 
Take the cross product of a vector v with itself and assign it to variable `cross_v`

```julia
## YOUR CODE GOES HERE
```

<!-- grade_id: cell-e6b6970ffe104df5, points: 1, grade: true, solution: false -->
```julia
@assert cross_v == [0, 0, 0]
```