# Basic linear algebra in Julia
Based on work by Andreas Noack Jensen (MIT) (https://github.com/andreasnoack)


First let's define a random matrix

```julia
A = rand(1:4, 3, 3)
```

Output:
```
3×3 Matrix{Int64}:
 1  4  2
 2  3  2
 4  1  1
```

Define a vector of ones

```julia
x = ones(3)
```

Output:
```
3-element Vector{Float64}:
 1.0
 1.0
 1.0
```

Many of the basic operations are the same as in other languages

```julia
b = A * x
```

Output:
```
3-element Vector{Float64}:
 7.0
 7.0
 6.0
```


#### Transposition

```julia
A'
```

Output:
```
3×3 adjoint(::Matrix{Int64}) with eltype Int64:
 1  2  4
 4  3  1
 2  2  1
```

and we can get the transpose with

```julia
transpose(A)
```

Output:
```
3×3 transpose(::Matrix{Int64}) with eltype Int64:
 1  2  4
 4  3  1
 2  2  1
```

#### Transposed multiplication
Julia allows us to write this without *

```julia
A'A
```

Output:
```
3×3 Matrix{Int64}:
 21  14  10
 14  26  15
 10  15   9
```

#### Solving linear systems 
The problem $Ax=b$ for ***square*** $A$ is solved by the \ operator.

```julia
x = A \ b
```

Output:
```
3-element Vector{Float64}:
 1.0
 0.9999999999999998
 1.0000000000000004
```

`A\b` gives us the *least squares solution* if we have an overdetermined linear system (a "tall" matrix)

```julia
Atall = rand(3, 2)
```

Output:
```
3×2 Matrix{Float64}:
 0.423735  0.375893
 0.747144  0.67883
 0.702932  0.285837
```

```julia
Atall \ b
```

Output:
```
2-element Vector{Float64}:
 6.632888991672504
 4.891429275931603
```

and the *minimum norm least squares solution* if we have a rank-deficient least squares problem

```julia
v = rand(3)
rankdef = hcat(v, v)
```

Output:
```
3×2 Matrix{Float64}:
 0.0246885  0.0246885
 0.37501    0.37501
 0.0918014  0.0918014
```

```julia
rankdef \ b
```

Output:
```
2-element Vector{Float64}:
 11.18696227587687
 11.186962275876876
```

Julia also gives us the minimum norm solution when we have an underdetermined solution (a "short" matrix)

```julia
bshort = rand(2)
Ashort = rand(2, 3)
```

Output:
```
2×3 Matrix{Float64}:
 0.262339  0.715286  0.71453
 0.600741  0.307078  0.0221058
```

```julia
Ashort \ bshort
```

Output:
```
3-element Vector{Float64}:
  0.3530108738618392
  0.04307411807582378
 -0.15361897630145296
```

# The LinearAlgebra library

While much of linear algebra is available in Julia by default (as shown above), there's a standard library named `LinearAlgebra` that brings in many more relevant names and functions. In particular, it provides factorizations and some structured matrix types.  As with all packages, you can bring these additional features into your session with a `using LinearAlgebra`.

### Exercises

#### 10.1 
Take the inner product (or "dot" product) of a vector `v` with itself and assign it to variable `dot_v`. 



```julia
v = [1, 2, 3]
v' * v
```

Output:
```
14
```

```julia
using LinearAlgebra
dot(v, v)
```

Output:
```
14
```

```julia
⋅ #\cdot + <<tab>>
dot_v = ⋅(v, v)
```

Output:
```
14
```

<!-- grade_id: cell-913fef9b0d19cd52, points: 1, grade: true, solution: false -->
```julia
@assert dot_v == 14
```

#### 10.2 
Take the outer product of a vector v with itself and assign it to variable `outer_v`

```julia
v = [1, 2, 3]
outer_v = v * v'
```

Output:
```
3×3 Matrix{Int64}:
 1  2  3
 2  4  6
 3  6  9
```

```julia
@assert outer_v == [1 2 3
                    2 4 6
                    3 6 9]
```

#### 10.3 
Take the cross product of a vector v with itself and assign it to variable `cross_v`

```julia
cross_v = cross(v, v)
```

Output:
```
3-element Vector{Int64}:
 0
 0
 0
```

<!-- grade_id: cell-e6b6970ffe104df5, points: 1, grade: true, solution: false -->
```julia
@assert cross_v == [0, 0, 0]
```