# Functions

Topics:
1. How to declare a function
2. Duck-typing in Julia
3. Mutating vs. non-mutating functions
4. Some higher order functions

## How to declare a function
Julia gives us a few different ways to write a function. The first requires the `function` and `end` keywords

```julia
function sayhi(name)
    println("Hi $name, it's great to see you!")
end
```

Output:
```
sayhi (generic function with 1 method)
```

```julia
sayhi("Class")
```

Output:
```
Hi Class, it's great to see you!
```

```julia
function f(x)
    x^2
end
```

Output:
```
f (generic function with 1 method)
```

Notice we didn't need to use a return statement. Julia defaults to the last used variable inside the function definition.

We can call either of these functions like this:

```julia
sayhi("C-3PO")
```

Output:
```
Hi C-3PO, it's great to see you!
```

```julia
f(42)
```

Output:
```
1764
```

```julia
function f(x)
    x^2
    y=10
end
```

Output:
```
f (generic function with 1 method)
```

What do we think the output of 
```julia
f(2)
```
will return?


```julia
f(2)
```

Output:
```
10
```

To overwrite this behaviour, we can use the `return` keyword

```julia
function f(x)
    return x^2
    y=10
end
```

Output:
```
f (generic function with 1 method)
```

```julia
f(2)
```

Output:
```
4
```

Alternatively, we could have declared either of these functions in a single line

```julia
sayhi2(name) = println("Hi $name, it's great to see you!")
```

Output:
```
sayhi2 (generic function with 1 method)
```

```julia
f2(x) = x^2
```

Output:
```
f2 (generic function with 1 method)
```

```julia
sayhi2("R2D2")
```

Output:
```
Hi R2D2, it's great to see you!
```

```julia
f2(42)
```

Output:
```
1764
```

Finally, we could have declared these as "anonymous" functions

```julia
sayhi3 = name -> println("Hi $name, it's great to see you!")
```

Output:
```
#1 (generic function with 1 method)
```

```julia
f3 = x -> x^2
```

Output:
```
#3 (generic function with 1 method)
```

```julia
sayhi3("Chewbacca")
```

Output:
```
Hi Chewbacca, it's great to see you!
```

```julia
f3(42)
```

Output:
```
1764
```

## Duck-typing in Julia
*"If it quacks like a duck, it's a duck."* <br><br>
Julia functions will just work on whatever inputs make sense. <br><br>
For example, `sayhi` works on the name of this minor tv character, written as an integer...

```julia
sayhi(55595472)
```

Output:
```
Hi 55595472, it's great to see you!
```

And `f` will work on a matrix. 

```julia
A = rand(3, 3)
A
```

Output:
```
3×3 Matrix{Float64}:
 0.448268   0.68648   0.147315
 0.0419332  0.477503  0.509792
 0.833519   0.34912   0.358262
```

```julia
f(A)
```

Output:
```
3×3 Matrix{Float64}:
 0.35252   0.686954  0.468776
 0.463742  0.434775  0.432244
 0.686897  0.863977  0.42912
```

`f` will also work on a string like "hi" because `*` is defined for string inputs as string concatenation.

```julia
f("hi")
```

Output:
```
"hihi"
```

On the other hand, `f` will not work on a vector. Unlike `A^2`, which is well-defined, the meaning of `v^2` for a vector, `v`, is not a well-defined algebraic operation. 

```julia
v = rand(3)
```

Output:
```
3-element Vector{Float64}:
 0.8587594861220774
 0.025955860869004366
 0.9668499457956325
```

```julia
f(v)
```

Output:
```
@ Main ./In[8]:2
 [3] top-level scope
   @ In[23]:1
```

We will see further examples of this behaviour when we cover **Multiple Dispatch**, a fundamental paradigm of the Julia Programming Language.

## Mutating vs. non-mutating functions

By convention, functions followed by `!` alter their contents and functions lacking `!` do not.

For example, let's look at the difference between `sort` and `sort!`.


```julia
v = [3, 5, 2]
```

Output:
```
3-element Vector{Int64}:
 3
 5
 2
```

```julia
sort(v)
```

Output:
```
3-element Vector{Int64}:
 2
 3
 5
```

```julia
v
```

Output:
```
3-element Vector{Int64}:
 3
 5
 2
```

`sort(v)` returns a sorted array that contains the same elements as `v`, but `v` is left unchanged. <br><br>

On the other hand, when we run `sort!(v)`, the contents of v are sorted within the array `v`.

```julia
sort!(v)
```

Output:
```
3-element Vector{Int64}:
 2
 3
 5
```

```julia
v
```

Output:
```
3-element Vector{Int64}:
 2
 3
 5
```

## Some higher order functions

### map

`map` is a "higher-order" function in Julia that *takes a function* as one of its input arguments. 
`map` then applies that function to every element of the data structure you pass it. For example, executing

```julia
map(f, [1, 2, 3])
```
will give you an output array where the function `f` has been applied to all elements of `[1, 2, 3]`
```julia
[f(1), f(2), f(3)]
```

```julia
map(f, [1, 2, 3])
```

Output:
```
3-element Vector{Int64}:
 1
 4
 9
```

Here we've squared all the elements of the vector `[1, 2, 3]`, rather than squaring the vector `[1, 2, 3]`.

To do this, we could have passed to `map` an anonymous function rather than a named function, such as

```julia
x -> x^3
```

Output:
```
#5 (generic function with 1 method)
```

via

```julia
map(x -> x^3, [1, 2, 3])
```

Output:
```
3-element Vector{Int64}:
  1
  8
 27
```

and now we've cubed all the elements of `[1, 2, 3]`!

### broadcast

`broadcast` is another higher-order function like `map`. `broadcast` is a generalization of `map`, so it can do every thing `map` can do and more. The syntax for calling `broadcast` is the same as for calling `map`

```julia
broadcast(f, [1, 2, 3])
```

Output:
```
3-element Vector{Int64}:
 1
 4
 9
```

and again, we've applied `f` (squared) to all the elements of `[1, 2, 3]` - this time by "broadcasting" `f`!

Some syntactic sugar for calling `broadcast` is to place a `.` between the name of the function you want to `broadcast` and its input arguments. For example,

```julia
broadcast(f, [1, 2, 3])
```
is the same as
```julia
f.([1, 2, 3])
```

```julia
f.([1, 2, 3])
```

Output:
```
3-element Vector{Int64}:
 1
 4
 9
```

Notice again how different this is from calling 
```julia
f([1, 2, 3])
```
We can square every element of a vector, but we can't square a vector!

To drive home the point, let's look at the difference between

```julia
f(A)
```
and
```julia
f.(A)
```
for a matrix `A`:

```julia
A = [i + 3*j for j in 0:2, i in 1:3]
```

Output:
```
3×3 Matrix{Int64}:
 1  2  3
 4  5  6
 7  8  9
```

```julia
f(A)
```

Output:
```
3×3 Matrix{Int64}:
  30   36   42
  66   81   96
 102  126  150
```

As before we see that for a matrix, `A`,
```
f(A) = A^2 = A * A
``` 

On the other hand,

```julia
B = f.(A)
```

Output:
```
3×3 Matrix{Int64}:
  1   4   9
 16  25  36
 49  64  81
```

contains the squares of all the entries of `A`.

This dot syntax for broadcasting allows us to write relatively complex compound elementwise expressions in a way that looks natural/closer to mathematical notation. For example, we can write

```julia
A .+ 2 .* f.(A) ./ A
```

Output:
```
3×3 Matrix{Float64}:
  3.0   6.0   9.0
 12.0  15.0  18.0
 21.0  24.0  27.0
```

instead of

```julia
broadcast(x -> x + 2 * f(x) / x, A)
```

Output:
```
3×3 Matrix{Float64}:
  3.0   6.0   9.0
 12.0  15.0  18.0
 21.0  24.0  27.0
```

and the two will perform exactly the same.

## Exercises

#### 2.1 
Write a function `add_one` that adds 1 to its input.

```julia

```

<!-- grade_id: cell-5119b9e9623c1cb7, points: 1, grade: true, solution: false -->
```julia
@assert add_one(1) == 2
```

Output:
```
Stacktrace:
 [1] top-level scope
   @ In[39]:1
```

<!-- grade_id: cell-50f83d27187a2064, points: 1, grade: true, solution: false -->
```julia
@assert add_one(11) == 12
```

Output:
```
Stacktrace:
 [1] top-level scope
   @ In[40]:1
```

#### 2.2 
Use `map` or `broadcast` to increment every element of matrix `A` by `1` and assign it to a variable `A1`.

```julia

```

```julia

```

```julia
@assert A1 == [2 3 4; 5 6 7; 8 9 10]
```

Output:
```
Stacktrace:
 [1] top-level scope
   @ In[41]:1
```

#### 2.3 
Use the broadcast dot syntax to increment every element of matrix `A1` by `1` and store it in variable `A2`

```julia

```

<!-- grade_id: cell-f3bd5479679a8fe1, points: 0, grade: true, solution: false -->
```julia
@assert A2 == [3 4 5; 6 7 8; 9 10 11]
```

Output:
```
Stacktrace:
 [1] top-level scope
   @ In[42]:1
```

```julia

```